# Assignment 0 

**Practical information**
* Handin no later than Sunday, **August 11th**.
* You must work alone for this assignment.
* We will host a workshop on **August 6th** where we cover the skills required to answer this problem set.
* If you have questions, ask them at the workshop or in a github issue on the sds repository.

This assignment is intended to ensure that you all know some basic python when you begin the summer school. We have included three different resources for learning python that you can use to learn the necessary skills for answering the questions. 

**A series of youtube videos:** [pythonprogramming.net fundamental](https://pythonprogramming.net/python-fundamental-tutorials/) (basics and intermediate)

**A tutorial website:** [The official python 3 tutorial](https://docs.python.org/3/tutorial/introduction.html) (sections 3, 4 and 5)

**A book:** [learn python the hard way](https://learnpythonthehardway.org/) 

You are free to pick from either of the above, or find your own ressources online. A very good (but perhaps slightly confusing) approach to learning anything programming related is to simply search google whenever you have a question, or cannot get your code to run. In 99% of the cases the first or second result will contain an answer to your question, and this saves you the time required to look for the right pages in a book, or the right tutorial on some website.


## A quick note on using this notebook
Below each problem you will find two code cells. The first one contains the line 
```python
raise NotImplementedError()
```
Erase this line and implement your question here. The cell below contains a number of `assert` statements. These check if your answer is correct. Once you have answered a question you can run the assert block. If you answer is correct nothing will happen, as the assertions all pass. If your answer is incorrect the assertions will raise an `AssertionError` and you will get a bit of information about the possible source of your mistake.

<br>
<br>
<br>
<br>
<br>
<br>

## Problem 0.1 - Build in types

Python has a number of build in _types_. These are basic data structures, such as numbers, sequences of letters (aka strings) and data containers which are conceptually similar to vectors, which you know from math. In this problem you will get a little bit of experience working with the different data types. 


### Problem 0.1.1

**Ex:** Add the two integers `9` and `2` and store the result in the variable named `answer_001`

<!-- grade_id: problem_001, points: 0, grade: false, solution: true -->
```python
# answer_001 = 

# YOUR CODE HERE
answer_001 = 9 + 2
```

<!-- grade_id: problem_001_tests, points: 1, grade: true, solution: false -->
```python
assert answer_001 == 11
```

### Problem 0.1.2
Above you added two integers together, and got a result of `11`. Python separates numbers in two classes, the _integers_ $...,-1,0,1,2,...$ and the _floats_, which are an approximation of the real numbers $\mathbb{R}$ (exactly how floats differ from reals is taught in introductory computer science courses).


**Ex:** 
* Add `4.2` to `2`, store the result in the variable `answer_0021`
* what is `0.33*10` in python? Store the result in the variable `answer_0022`

> *Note:* the notebook will only print out the result of the last line in each cell.

<!-- grade_id: problem_002, points: 0, grade: false, solution: true -->
```python
# answer_0021 = 
# answer_0022 =

# YOUR CODE HERE
answer_0021 = 4.2 + 2
answer_0022 = 0.33*10

```

<!-- grade_id: problem_002_tests, points: 1, grade: true, solution: false -->
```python
assert answer_0021 == 6.2
assert answer_0022 == 0.33*10
assert type(answer_0022) == float
```

### Problem 0.1.3
Python also has a build in data type called a _string_. This is simply a sequence of letters/characters and can thus contain names, sentences etc. To define a string in python, you need to wrap your sentence in either double or single quotation marks. For example you could write `"Hello world!"`. 

**Ex:**  In python the `+` is not only used for numbers. Use the `+` to add together the three strings `"Social"`, `"Data"` and `"Science"`. What is the result? Store it in a new variable called `answer_003`

<!-- grade_id: problem_003, points: 0, grade: false, solution: true -->
```python
# answer_003 = 

# YOUR CODE HERE
answer_003 = 'Social' + 'Data' + 'Science'
```

<!-- grade_id: problem_003_tests, points: 1, grade: true, solution: false -->
```python
assert answer_003 == 'SocialDataScience'
```

### Problem 0.1.4
Now that you have seen all of the _elementary data types_ we will look at some _compound data types_. There are three of these, the _tuple_, the _list_ and the _dictionary_. 

* The _tuple_ is declared with round parentheses, e.g. `(1,2,3)` each element in the tuple is separated by a comma. One you have declared a tuple you cannot change it's content without making a copy of the tuple first (you will read that the tuple is an _immutable_ data type).


* The _list_ is almost identical to the tuple. It is declared using square parentheses, e.g. `[1,2,3]`. Unlike the tuple, a list can be changed after definition, by adding, changing or removing elements. This is called a _mutable_ data type.


* The _dictionary_ or simply _dict_ is also a mutable data type. Unlike the other data types the dict works like a lookup table, where each element of data stored in the dictionary is associated with a name. To look up an item in the dictionary you don't need to know its position in the dictionary, only its name. The dict is defined with curly braces and a colon to separate the name from the value, e.g. `{'name_of_first_entry': 1, 'name_of_second_entry: 2}`.

**Ex:** Define the variable `y` as a list containing the elements `0, 4, 'a', 'g'`. Also define a variable `z` which is a tuple containing the same elements. Try to access the 0th element of `y` and the 1st element of `z`. Store these two values in the two variables `answer_004_y0` and `answer_004_z1`.

> *Hint:* To access the _n_'th element of a list/tuple write `myData[n]`, for example `y[0]` gets the 0th element of `y`. 

<!-- grade_id: problem_004, points: 0, grade: false, solution: true -->
```python
# answer_004_y0 = 
# answer_004_z1 = 

# YOUR CODE HERE
y = [0, 4, 'a', 'g']
z = (0,4,'a','g')
answer_004_y0 = y[0]
answer_004_z1 = z[1]
```

<!-- grade_id: problem_004_tests, points: 1, grade: true, solution: false -->
```python
assert y == [0, 4, 'a', 'g']
assert z == (0, 4, 'a', 'g')
assert answer_004_y0 == 0
assert answer_004_z1 == 4
```

<br>
<br>
<br>
<br>
<br>
<br>

## Problem 0.2 - Control flow
_Control flow_ means writing code that controls the way data or information flows through the program. The concepts of control flow should be recognizable outside of coding as well. For example when you go shopping you might want to buy koldskål, but only **if** the kammerjunker are on sale. **else** you will buy icecream. These kinds of logic come up everywhere in coding; self driving cars should go forward only **if** the light is green, items should be listed for sale in a web shop only **if** they are in stock, stars should be put on the estimates **if** they are significant etc. 

Another kind of control flow deals with doing things repeatedly. For example dishes should be done **while** there are still dirty dishes to wash, **for** each student **in** social data science a grade should be given, etc.

In this problem you will work with both kinds of control flow. You will not learn about `break` and `continue`, or about `try` and `except` which are also control flow statements. These are slightly more advanced, but it can be a good idea to look them up yourself. 

### Problem 0.2.1
In python the if/else logic consists of three keywords: *if*, *elif* (else if) and *else*. The if and elif keywords should be followed by a logical statement that is either `True` or `False`. The code that should be executed `if` the logic is `True` is written on the lines below the `if`, and should be indented one TAB (or 4 spaces). Finally all control flow statements must end with a colon. 

**Ex:** Read the code in the cell below. When you run it you will be asked to supply a value of `x`. Supply a value of `x` that makes the code print "Good job!"

<!-- grade_id: problem_021, points: 0, grade: false, solution: true -->
```python
x = float(input("Please supply a value for x "))

if x < 5:
    answer_021_sentinel = 0
    print("x is too small")

elif x >= 5 and x < 8:
    answer_021_sentinel = 1
    print("Good job!")

else: 
    answer_021_sentinel = 0    
    print("x is too large")
    
    
# YOUR CODE HERE
```

<!-- grade_id: problem_021_tests, points: 1, grade: true, solution: false -->
```python
assert answer_021_sentinel == 1
```

### Problem 0.2.2
Above we used two different types of comparison: `>=` and `<`. To compare two values and check if they are equal python uses double equal signs `==` (remember a single = was used to assign values to a variable). 

**Ex:** The code below draws a random number between 0 and 1 and stores in the variable `randnum`. Write an if/else block that defines a new variable `answer_022` which is equal to 0 if `randnum <= 0.3` and is 1 if `randnum > 0.3`. That is write code that defines 

<br>
$$
\text{answer_022} = \begin{cases}
0 \text{  if } \text{randnum} \leq 0.3 \\
1 \text{  if } \text{randnum} > 0.3
\end{cases}
$$

<!-- grade_id: problem_022, points: 0, grade: false, solution: true -->
```python
import random
randnum = random.uniform(0,1)

# YOUR CODE HERE
if randnum <= 0.3:
    answer_022 = 0
if randnum > 0.3:
    answer_022 = 1
else:
    print("Out of bounds")
```

<!-- grade_id: problem_022_tests, points: 1, grade: true, solution: false -->
```python
assert (answer_022 == 0 and randnum <= 0.3) or (answer_022 == 1 and randnum > 0.3)
```

### Problem 0.2.3

control flow that does the same thing repeatedly is called a _loop_. In python you can loop through anything that is _iterable_, e.g. anything where it makes sense to say "for each element in this item, do whatever." Lists, tuples and dictionaries are all iterable, and can thus be looped over. This kind of loop is called a _for loop_. The basic syntax is 

```python
for element in SomeIterable:
    dosomething(element)
```
where `element` is a temporary name given to the current element reached in the loop, and `dosomething()` can be any valid python code. 

**Ex:** Begin by initializing an emply list in the variable `answer_023` (simply write `answer_023=[]`). Then loop trough the list `y` that you defined in problem 1. For each element in `y`, append 3 times the element to `answer_023`. (You can finish off by showing the content of `answer_023` after the loop has run.)

> *Hint:* To append data to a list you can write `answer_023.append(new_element)` where `new_element` is the new data that you want to append.

<!-- grade_id: problem_023, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
answer_023 = []
iterations = 3
for item in y:
    current = item
    iterations_performed = 1
    while iterations > iterations_performed:
        current = current + item
        iterations_performed = iterations_performed + 1
    answer_023.append(current)
```

<!-- grade_id: problem_023_tests, points: 1, grade: true, solution: false -->
```python
assert answer_023 == [0, 12, 'aaa', 'ggg']
```

### Problem 0.2.4
There is only two kinds of loops in python. In addition to the for loop there is the _while loop_. Instead of looping over an iterable the `while` loop continues going as long as a supplied logical condition is True. Most commonly the while loop is combined with a counting variable that keeps track of how many times the loop has been run. The basic syntax is

```python 

i = 0
while i < 100:
    dosomething()    
    i = i + 1
```
This code will run 100 times before stopping. At each iteration it checks that `i` is smaller than 100. If it is, it does something and adds 1 to the variable `i` before repeating.

**Ex:** Begin by defining an empty list named `answer_024`. Write a while loop that runs 1000 times from $i=0$ to $i=1000$. In each loop it should determine if the current value of `i` is a multiple of 27. If it is, append the number to the answer list `answer_024`. (recall that $i$ is divisible by $a$ if $i \text{ mod } a = 0$. The modulo operator in python is `%`)
> *Hint:* The `if` statement does not need to be followed by an `else`. You can simply code the `if` part and python will automatically skip it and continue if the logical condition is False.
>
> *Hint:* Remember to increment `i` in each iteration. Otherwise the program will run forever. If this happens press _kernel > interrupt_ in the menu bar. 

<!-- grade_id: problem_024, points: 0, grade: false, solution: true -->
```python
i = 0
answer_024 = []

# YOUR CODE HERE
while i < 1000:
    if i%27 == 0:
        answer_024.append(i)
    i = i +1

```

<!-- grade_id: problem_024_tests, points: 1, grade: true, solution: false -->
```python
assert answer_024 == [i for i in range(1000) if i % 27 == 0]
```

<br>
<br>
<br>
<br>
<br>
<br>

## Problem 0.3 - functions
If you have never programmed in anything but statistical software such as Stata or SAS the concept of functions might be new to you. In python a function is simply a "recipe" that is first written, and then later used to compute something. Conceptually functions in programming are similar to functions in math. They have between $0$ and "$\infty$" inputs, do some calculation using their inputs and then return between 1 and "$\infty$" outputs. 

Python contains a large number of build-in functions. Below you are given examples of how to use the most commonly used build-ins. You should make yourself comfortable using each of the functions shown below.

```python
# Setup for the examples. We define two lists
# to show you the build in functions.
l1 = [0,10,20,30,40,50,60,70,80,90]
l2 = ['a', 'b', 'c', 'd', 'e', 'f','g', 'h', 'i', 'j']
```

```python
# The len(x) function gives you the length of the input
len(l2)
```

```python
# The abs(x) function returns the absolute value of x
abs(-5)
```

```python
# The min(x) and max(x) functions return the minimum
# and maximum of the input.
min(l1), max(l1)
```

```python
# The map(function, Iterable) function applies the 
# supplied function to each element in Iterable:
# (the list() call just converts the result to a list)
list(
    map(len, l2)
)
```

```python
# The range([start], stop, [step]) function returns 
# a range of numbers from `start` to `stop`, 
# in increments of `step`. The values in [] are 
# optional. If no start value is set, it defaults to 0.
# If no step value is set it defaults to 1. A stop value
# must always be set.

print("Range from 0 to 100, step=1:", range(100))
print("Range from 0 to 100, step=2", range(0, 100, 2))
print("Range from 10 to 65, step=3", range(10, 65, 3))
```

```python
# The reversed(x) function reverses the input.
# we can then loop trough it backwards
l1_reverse = reversed(l1)

for e in l1_reverse:
    print(e)
```

```python
# The enumerate(x) function returns the 
# index of the item as well as the item 
# itself in sequence. With it you can 
# loop through things while keeping track
# of their position:
l2_enumerate = enumerate(l2)

for index, element in l2_enumerate:
    print(index, element)
```

```python
# The zip(x,y,...) function "zips" together two
# or more iterables allowing you to loop through
# them pairwise:

l1l2_zip = zip(l1,l2)

for e1, e2 in l1l2_zip:
    print(e1, e2)

```

### Problem 0.3.1
You can also write your own python functions. A python function is defined with the `def` keyword, followed by a user defined name of the function, the inputs to the function and a colon. On the following line the _function body_ is written, indented by one TAB. 

Functions use the keyword `return` to signal what values the function should return after doing its calculations on the inputs. 

For example we can define a function named `my_first_function` that takes two inputs `a` and `b`, computes the values $a^b$ (written in python as `a**b`) and $b^a$ and returns the larger of the two, in the following way:

```python
def my_first_function(a,b):
    v1 = a**b
    v2 = b**a
    
    if v1 > v2:
        return v1
    else:
        return v2
```

**Ex:** Write a function called `average` that takes as input a list of numbers, and returns the average of the values in the list. Use your function to calculate the average of the values `[1,2,3,4,5,6,7,8,9]`. Store the result of this computation in a variable named `answer_031`.

<!-- grade_id: problem_031, points: 0, grade: false, solution: true -->
```python
# answer_031 = 

# YOUR CODE HERE
def average(numbers):
    z_sum = 0
    for element in numbers:
        z_sum = z_sum + element
    z_average = z_sum / len(numbers)
    return z_average
answer_031 = average([1,2,3,4,5,6,7,8,9])
```

<!-- grade_id: problem_031_tests, points: 1, grade: true, solution: false -->
```python
assert answer_031 == 5
assert average([1,1]) == 1
assert average([-2,2]) == 0
assert average([1]) == 1
```

### Problem 0.3.2
Recall that [Eulers constant](https://en.wikipedia.org/wiki/E_(mathematical_constant)) $e$ can be calculated as 
$$
e=\lim_{n\rightarrow \infty}\left(1+\frac{x}{n}\right)^{n}
$$
of course we cannot compute the limit on a finite memory computer. Instead we can calculate approximations by taking $n$ large enough.

**Ex:** write a function named `eulers_e` that takes two inputs `x` and `n`, calculates 
$$
\left(1+\frac{x}{n}\right)^{n}
$$
and returns this value. Use your function to calculate `eulers_e(1, 50)` and store this value in the variable `answer_032`.

<!-- grade_id: problem_032, points: 0, grade: false, solution: true -->
```python
# answer_032 = 

# YOUR CODE HERE
def eulers_e(x,n):
    return (1+(x/n))**n
answer_032 = eulers_e(1,50)
    
```

<!-- grade_id: problem_032_tests, points: 1, grade: true, solution: false -->
```python
assert round(answer_032, 2) == 2.69
assert eulers_e(1,1) == 2
assert eulers_e(0,19) == 1
assert round(eulers_e(1,1000), 2) == 2.72
```

### Problem 0.3.3
The inverse of the exponential is the [logarithm](https://en.wikipedia.org/wiki/Logarithm). Like the exponential function there are limit definitions of the logarithm. One of these is 
$$
log(x) = 2 \cdot \sum_{k=0}^{\infty} \frac{1}{2k+1} \left( \frac{x-1}{x+1} \right)^{2k+1}
$$

where $\sum_{k=0}^{\infty}$ signifies the sum of infinitely many elements, starting from $k=0$. Each element in the sum takes the value $\frac{1}{2k+1} \left( \frac{x-1}{x+1} \right)^{2k+1}$ for some $k$. As before we mus approximate this with a finite sum. 

**Ex:** define another function called `natural_logarithm` which takes two inputs `x` and `k_max`. In the function body calculate 
$$
2 \cdot \sum_{k=0}^{k\_max} \frac{1}{2k+1} \left( \frac{x-1}{x+1} \right)^{2k+1}
$$
and return this value. Use your function to calculate `natural_logarithm(2.71, 1000)`. Store this value in a variable named `answer_033`.

> *Hint:* to calculate the sum, first initialize a value total = 0, and loop through $k\in [0, k\_max]$. Add the partial sum to your total in each step of the loop. After finalizing the loop you can then multiply the total by 2 and return the result.

<!-- grade_id: problem_033, points: 0, grade: false, solution: true -->
```python
# answer_033 = 

# YOUR CODE HERE
def natural_logarithm(x, k_max):
    sum_so_far = 0
    k = 0
    while k<k_max:
        sum_so_far = sum_so_far + (1/(2*k+1))*(((x-1)/(x+1))**(2*k+1))
        k = k+1
    return (sum_so_far * 2)
answer_033 = natural_logarithm(2.71, 1000)
```

<!-- grade_id: problem_033_tests, points: 1, grade: true, solution: false -->
```python
assert round(answer_033) == 1
assert natural_logarithm(1,1) == 0
assert natural_logarithm(0,1) == -2
assert natural_logarithm(1,100) == 0
```

### Problem 0.3.4
Just like numbers, strings and data types, python treats functions as an object. This means you can write functions, that take a function as input and functions that return functions after being executed. This is sometimes a useful tool to have when you need to add extra functionality to an already existing function, or if you need to write _function factories_.

**Ex:** Write a function called `exponentiate` that takes one input named `func`. In the body of `exponentiate` define a nested function (i.e. a function within a function) called `with_exp` that takes two inputs `x` and `k`. The nested function should return `func(e, k)` where `e=eulers_e(x, k)`. The outer function should return the nested function `with_exp`, i.e. write something like

```python
def exponentiate(func):
    def with_exp(x, k):
        e = eulers_e(x, k)
        value = #[FILL IN]
        return value
    return with_exp
```

Call the `exponentiate` function on `natural_logarithm` and store the result in a new variable called `logexp`. Then call `logexp(1, 1000)` and store this value in the variable `answer_034`.

<!-- grade_id: problem_034, points: 0, grade: false, solution: true -->
```python
# logexp = 
# answer_034 = 

# YOUR CODE HERE
def exponentiate(func):
    def with_exp(x, k):
        e = eulers_e(x, k)
        value = func(e, k)
        return value
    return with_exp
logexp = exponentiate(natural_logarithm)
answer_034 = logexp(1,1000)
```

<!-- grade_id: cell-4d43e938f5e89487, points: 1, grade: true, solution: false -->
```python
assert round(answer_034) == 1
assert round(logexp(2, 1000)) == 2
assert round(logexp(3, 1000)) == 3
assert round(logexp(4, 1000)) == 4
```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```

```python

```