# Mandatory Assignment 1

This is the first of two mandatory assignments which must be completed during the course. First some practical information:

* When is the assignment due?: **23:59, Sunday, August 19, 2018.**
* How do you grade the assignment?: You will **peergrade** each other as primary grading. 
* Can i work with my group?: **yes**

The assigment consist of one to tree problems from each of the exercise sets you have solved so far (excluding Set 1). We've tried to select problems which are self contained, but it might be nessecary to solve some of the previous exercises in each set to fully answer the problems in this assignment.

## Problems from Exercise Set 2:

> **Ex. 2.2**: Make two lists. The first should be numbered. The second should be unnumbered and contain at least one sublevel. 

```python
# [Answer to Ex. 2.2 here] (convert to markdown cell)
```

## Problems from Exercise set 3:

> **Ex. 3.1.3:** Let `l1 = ['r ', 'Is', '>', ' < ', 'g ', '?']`. Create from `l1` the sentence `"Is r > g?"` using your knowledge about string formatting. Store this new string in a variable called `answer_31`. Make sure there is only one space in between worlds.
>
>> _Hint:_ You should be able to combine the above informations to solve this exercise.

<!-- grade_id: problem_31, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.1.3 here]
l1 = ['r ', 'Is', '>', ' < ', 'g ', '?']

# answer_31 = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_31_tests, points: 1, grade: true, solution: false -->
```python
assert answer_31 == "Is r > g?"
```

> **Ex. 3.1.4**: Create an empty dictionary `words` using the `dict()`function. Then add each of the words in `['animal', 'coffee', 'python', 'unit', 'knowledge', 'tread', 'arise']` as a key, with the value being a boolean indicator for whether the word begins with a vowel. The results should look like `{'bacon': False, 'asynchronous': True ...}`. Store the result in a new variable called `answer_32`.
>
>> _Hint:_ You might want co first construct a function that asseses whether a given word begins with a vowel or not.

<!-- grade_id: problem_32, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.1.4 here]
W = ['animal', 'coffee', 'python', 'unit', 'knowledge', 'tread', 'arise']

# answer_32 = 


# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: problem_32_tests, points: 1, grade: true, solution: false -->
```python
assert answer_32 == {i: i[0] in 'aeiou' for i in W}
assert sorted(answer_32) == sorted(W)
```

> **Ex. 3.3.2:** use the `requests` module (get it with `pip install requests`) and `construct_link()` which you defined in the previous question (ex 3.3.1) to request birth data from the "FOD" table. Get all available years (variable "Tid"), but only female births (BARNKON=P) . Unpack the json payload and store the result. Wrap the whole thing in a function which takes an url as input and returns the corresponding output.
>
> Store the birth data in a new variable called `answer_33`.
>
>> _Hint:_ The `requests.response` object has a `.json()` method. 
>
>> _Note:_ you wrote `construct_link()` in 3.3.1, if you didn't heres the link you need to get: `https://api.statbank.dk/v1/data/FOLK1A/JSONSTAT?lang=en&Tid=*`

<!-- grade_id: problem_33, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 3.3.2 here]

# answer_33 = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_33_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(answer_33['dataset'].keys()) == ['dimension', 'label', 'source', 'updated', 'value']
assert 'BARNKON' in answer_33['dataset']['dimension'].keys()
```

## Problems from exercise set 4


```python
import numpy as np
import pandas as pd
```

> **Ex. 4.1.1:** Use Pandas' CSV reader to fetch  daily data weather from 1864 for various stations - available [here](https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/by_year/). Store the dataframe in a variable called `answer_41`.
>
>> *Hint 1*: for compressed files you may need to specify the keyword `compression`.
>
>> *Hint 2*: keyword `header` can be specified as the CSV has no column names.
>
>> *Hint 3*: Specify the path, as the URL linking directly to the 1864 file. 

<!-- grade_id: problem_41, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.1 here]

# answer_41 = 

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: problem_41_tests, points: 1, grade: true, solution: false -->
```python
assert answer_41.shape == (27349, 8)
assert list(answer_41.columns) == list(range(8))
```

> **Ex. 4.1.2:** Structure your weather DataFrame by using only the relevant columns (station identifier, data, observation type, observation value), rename them. Make sure observations are correctly formated (how many decimals should we add? one?).
>
> Store the resulting dataframe in a new variable called `answer_42`.
>
>> *Hint:* rename can be done with `df.columns=COLS` where `COLS` is a list of column names.

<!-- grade_id: problem_42, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.2 here]

# answer_42 = 

# YOUR CODE HERE
raise NotImplementedError()


```

<!-- grade_id: problem_42_tests, points: 1, grade: true, solution: false -->
```python
assert answer_42.shape == (27349, 4)
assert 144.8 in [answer_42[i].max() for i in answer_42]
assert -666.0 in [answer_42[i].min() for i in answer_42]
assert 18640101 in [answer_42[i].min() for i in answer_42]
```

> **Ex. 4.1.3:**  Select data for the station `ITE00100550` and only observations for maximal temperature. Make a copy of the DataFrame. Explain in a one or two sentences how copying works.
>
> Store the subsetted dataframe in a new variable called `answer_43`.
>
>> *Hint 1*: the `&` operator works elementwise on boolean series (like `and` in core python).
>
>> *Hint 2*: copying of the dataframe is done with the `copy` method for DataFrames.

<!-- grade_id: problem_43, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.3 here]

# answer_43 = 


# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_43_tests, points: 1, grade: true, solution: false -->
```python
assert 'ITE00100550' in [answer_43[i].min() for i in answer_43]
assert 'ITE00100550' in [answer_43[i].max() for i in answer_43]
assert 'TMAX' in [answer_43[i].min() for i in answer_43]
assert 'TMAX' in [answer_43[i].max() for i in answer_43]
```

> **Ex. 4.1.4:** Make a new column in `answer_44` called `TMAX_F` where you have converted the temperature variables to Fahrenheit. Make sure not to overwrite `answer_43`.
>
> Store the resulting dataframe in a variable called `answer_44`.
>
>> *Hint*: Conversion is $F = 32 + 1.8*C$ where $F$ is Fahrenheit and $C$ is Celsius.

<!-- grade_id: problem_44, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 4.1.4 here]
answer_44 = answer_43.copy()
# answer_44 = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_44_tests, points: 1, grade: true, solution: false -->
```python
assert set(answer_44.columns) -  set(answer_43.columns) == {'TMAX_F'}
```

## Problems from exercise set 5

```python
import matplotlib.pyplot as plt
import numpy as np 
import pandas as pd
import seaborn as sns 

%matplotlib inline 

iris = sns.load_dataset('iris')
titanic = sns.load_dataset('titanic')
```

> **Ex. 5.1.1:**: Show the first five rows of the titanic dataset. What information is in the dataset? Use a barplot to show the probability of survival for men and women within each passenger class. Can you make a boxplot showing the same information (why/why not?). _Bonus:_ show a boxplot for the fare-prices within each passenger class. 
>
> Spend five minutes discussing what you can learn about the survival-selection aboard titanic from the figure(s).
>
> > _Hint:_ https://seaborn.pydata.org/generated/seaborn.barplot.html, specifically the `hue` option.


<!-- grade_id: problem_51, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.1 here]

# YOUR CODE HERE
raise NotImplementedError()
```

> **Ex. 5.1.2:** Using the iris flower dataset, draw a scatterplot of sepal length and petal length. Include a second order polynomial fitted to the data. Add a title to the plot and rename the axis labels.
> _Discuss:_ Is this a meaningful way to display the data? What could we do differently?
>
> For a better understanding of the dataset this image might be useful:
> <img src="iris_pic.png" alt="Drawing" style="width: 200px;"/>
>
>> _Hint:_ use the `.regplot` method from seaborn. 

<!-- grade_id: cell-8631e73e3796008b, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.2 here]

# YOUR CODE HERE
raise NotImplementedError()
```

> **Ex. 5.1.3:** Combine the two of the figures you created above into a two-panel figure similar to the one shown here:
> <img src="Example.png" alt="Drawing" style="width: 600px;"/>
>
> Save the figure as a png file on your computer. 
>> _Hint:_ See [this question](https://stackoverflow.com/questions/41384040/subplot-for-seaborn-boxplot) on stackoverflow for inspiration.

<!-- grade_id: cell-686e1fcbc475435a, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.3 here]

# YOUR CODE HERE
raise NotImplementedError()
```

> **Ex. 5.1.4:** Use [pairplot with hue](https://seaborn.pydata.org/generated/seaborn.pairplot.html) to create a figure that clearly shows how the different species vary across measurements. Change the color palette and remove the shading from the density plots. _Bonus:_ Try to explain how the `diag_kws` argument works (_hint:_ [read here](https://stackoverflow.com/questions/1769403/understanding-kwargs-in-python))

<!-- grade_id: cell-0f5793c056a02040, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 5.1.4 here]

# YOUR CODE HERE
raise NotImplementedError()
```

## Problems from exercise set 6

> _Note:_ In the exercises we asked you to download weather data from the NOAA website. For this assignment the data are loaded in the following code cell into two pandas dataframes.

```python
%matplotlib inline
import pandas as pd
import matplotlib.pyplot as plt

weather_1864 = pd.read_csv('weather_data_1864.csv')
```

> **Ex. 6.1.4:** Extract the country code from the station name into a separate column.
>
> Create a new column in `weather_1864` called `answer_61` and store the country codes here.
>
>> _Hint:_ The station column contains a GHCND ID, given to each weather station by NOAA. The format of these ID's is a 2-3 letter country code, followed by a integer identifying the specific station. A simple approach is to assume a fixed length of the country ID. A more complex way would be to use the [`re`](https://docs.python.org/2/library/re.html) module.

<!-- grade_id: problem_61, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 6.1.4]
# weather_1864['answer_61'] =

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_61_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(weather_1864['answer_61'].str[:2].unique()) == sorted(['SZ', 'CA', 'EZ', 'GM', 'AU', 'IT', 'BE', 'UK', 'EI', 'AG', 'AS'])
```

> **Ex. 6.1.5:** Make a function that downloads and formats the weather data according to previous exercises in Exercise Section 4.1, 6.1. You should use data for ALL stations but still only select maximal temperature. _Bonus:_ To validate that your function works plot the temperature curve for each country in the same window. Use `plt.legend()` to add a legend. 
>
> Name your function `prepareWeatherData`.

<!-- grade_id: problem_62, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 6.1.5]

def prepareWeatherData(year):
    # Your code here
    return 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_62_tests, points: 1, grade: true, solution: false -->
```python
assert prepareWeatherData('1864').shape == (5686, 6)
```

## Problems from exercise set 7

> _Note:_ Once again if you haven't managed to download the data from NOAA, you can refer to the github repo to get csv-files containing the required data.

```python
%matplotlib inline 

import pandas as pd 
import matplotlib.pyplot as plt
import matplotlib as mpl

# Increases the plot size a little
mpl.rcParams['figure.figsize'] = 11, 6
```

> **Ex. 7.1.1:** Plot the monthly max,min, mean, first and third quartiles for maximum temperature for our station with the ID _'ITE00100550'_ in 1864. 

> *Hint*: the method `describe` computes all these measures.

<!-- grade_id: problem_71, points: 1, grade: true, solution: true -->
```python
# [Answer to Ex. 7.1.1]

# YOUR CODE HERE
raise NotImplementedError()
```

> **Ex. 7.1.2:** Get the processed data from years 1864-1867 as a list of DataFrames. Convert the list into a single DataFrame by concatenating vertically. 
>
> Name the concatenated data `answer_72`

<!-- grade_id: problem_72, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 7.1.2]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_72_tests, points: 1, grade: true, solution: false -->
```python
assert answer_72.shape == (30003, 7)
```

> **Ex. 7.1.3:** Parse the station location data which you can find at https://www1.ncdc.noaa.gov/pub/data/ghcn/daily/ghcnd-stations.txt. Merge station locations onto the weather data spanning 1864-1867.  
>
> Store the merged data in a new variable called `answer_73`.
>
> _Hint:_ The location data have the folllowing format, 

```
------------------------------
Variable   Columns   Type
------------------------------
ID            1-11   Character
LATITUDE     13-20   Real
LONGITUDE    22-30   Real
ELEVATION    32-37   Real
STATE        39-40   Character
NAME         42-71   Character
GSN FLAG     73-75   Character
HCN/CRN FLAG 77-79   Character
WMO ID       81-85   Character
------------------------------
```

> *Hint*: The station information has fixed width format - does there exist a pandas reader for that?

<!-- grade_id: problem_73, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 7.1.3]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_73_tests, points: 1, grade: true, solution: false -->
```python
assert answer_73.shape == (5686, 15) or answer_73.shape == (30003, 15)
```

## Problems from exercise set 8

> **Ex. 8.1.2.:** Use the `request` module to collect the first page of job postings.
>
> Store the response.json() object in a new variable called `answer_81`.
> 

<!-- grade_id: problem_81, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 8.1.2]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_81_tests, points: 1, grade: true, solution: false -->
```python
assert sorted(answer_81.keys()) == sorted(['Expression', 'Facets', 'JobPositionPostings', 'TotalResultCount'])
```

> **Ex. 8.1.3.:** Store the 'TotalResultCount' value for later use. Also create a dataframe from the 'JobPositionPostings' field in the json. Name this dataframe `answer_82`.

<!-- grade_id: problem_82, points: 0, grade: false, solution: true -->
```python
# [Answer to Ex. 8.1.3]
# answer_82 = 

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: problem_82_tests, points: 1, grade: true, solution: false -->
```python
assert answer_82.shape == (20,44)
```